Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select  **Restart** and **Clear All Outputs**, then select **Run All**).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Jackson Chang"
COLLABORATORS = "GENOS"
```

---

## **Problem Set 3** Execution Control I

This is a problem set. You can work on it with the help of your classmates, Google, or ask me any questions you would like.

I encourage discussion during these problem sets.  

```python
import numpy as np
```

1. Write a function called **clip** that takes as input arguments an array of integers **x** and two values **upper** and **lower**.  The function replaces any values in the array greater than **upper** with the value **upper** and lower thant **lower** with the value **lower**.  The function should return the new array with modified values.  Solve the problem using logical indexing.       

<!-- grade_id: cell-ed628adef9122d0e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def clip(x, upper, lower):
    x = np.array(x)
    x[x>upper] = upper
    x[x<lower] = lower
    return x
    return lower
    return upper
```

<!-- grade_id: cell-338ae9bd6b6478af, points: 2, grade: true, solution: false -->
```python
x = np.arange(10,200,10)
upper = 100
lower = 50
x = clip(x, upper, lower)
assert x[3] == 50  
assert x[10] == 100      
```

2.  A store is having a sale based on what the customer spends. 

* If they spend less than 100, they receive no discount.
* If the user spends between 100 and 150, they receive 25 off.
* If they user spends 150 or more, they receive 50 off. 

Write a function **discount** that calculates the sale price the customer pays.

The argument for this function is the price of the items the customer purchases. 

<!-- grade_id: cell-59ac799a002e3ebb, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def discount(x):
    if x>=150:
        return x-50
    elif (100<=x) & (x<150):
        return x-25
    else:
       return x
```

<!-- grade_id: cell-15bef6706414f9c7, points: 2, grade: true, solution: false -->
```python
price = 90
saleprice = discount(price)
assert saleprice == price
```

<!-- grade_id: cell-3d940bd592341be5, points: 2, grade: true, solution: false -->
```python
price = 120
saleprice = discount(price)
assert saleprice == price-25
```

<!-- grade_id: cell-d02ff3e5ae414724, points: 2, grade: true, solution: false -->
```python
price = 150
saleprice = discount(price)
assert saleprice == price-50
```

3. Write a function called **trianglename**. 
- The input arguments to the function are the lengths of the 3 sides of the triangle - a, b and c.  
- Your function should use a statement of the type (`if`-`elif`-`else`) to determine if the triangle is 'equilateral', 'isosceles' or 'obtuse' and return the correct triangle name 
* 'equilateral' means all sides are equal, 
* 'isosceles' means two of the sides are equal but not the third one, 
* 'obtuse' means all 3 sides are different.
- There is no preferred way to do this problem, and there are many possible solutions. 
- Note that I have created 3 test cases here. When testing a piece of code, you should always test that all possible cases work. 

<!-- grade_id: cell-6a863cd8f2f710ce, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def trianglename(a,b,c):
    if a == b == c:
        return "equilateral"
    elif (a == b != c)| (a != b == c) | (a == c !=b):
        return 'isosceles'
    else:
        return 'obtuse' 
```

<!-- grade_id: cell-616f43338754cde7, points: 2, grade: true, solution: false -->
```python
x = 6 
y = 7
z = 7
triangle = trianglename(x,y,z)
assert triangle == 'isosceles'
```

<!-- grade_id: cell-bf336b105b6ecee5, points: 2, grade: true, solution: false -->
```python
x = 7 
y = 7
z = 7
triangle = trianglename(x,y,z)
assert triangle == 'equilateral'
```

<!-- grade_id: cell-d88d286035d01233, points: 2, grade: true, solution: false -->
```python
x = 6 
y = 7
z = 8
triangle = trianglename(x,y,z)
assert triangle == 'obtuse'
```

4.  Write a function  that checks if a positive whole number provided as an argument is a prime number. The function should return the string 'prime' if the number is prime, and 'not prime' if the number is not prime. A prime number is divisible by exactly two numbers: itself and 1.  Do not use any loops (if you know how to use them).

   Hint:  
* % tells you the remainder when dividing a number by a number or an array
* np.arange can be used to make an array

<!-- grade_id: cell-abac5d0d0749ddd6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def isprime(n):
    if n<2:
        return 'not prime'
    division = np.arange(1, n+1)
    remainder_zero = n%division == 0
    if np.sum(remainder_zero) == 2:
        return 'prime'
    else:
        return 'not prime'
```

<!-- grade_id: cell-735b98f4df1eca7d, points: 2, grade: true, solution: false -->
```python
assert isprime(8) == 'not prime'
### BEGIN HIDDEN TESTS ###
assert isprime(7) == 'prime'
### END HIDDEN TESTS ###
```

<!-- grade_id: cell-13196bb4dcd60f06, points: 2, grade: true, solution: false -->
```python
assert isprime(31) == 'prime'
### BEGIN HIDDEN TESTS ###
assert isprime(30) == 'not prime'
### END HIDDEN TESTS ###
```

<!-- grade_id: cell-634914584a356eb3, points: 2, grade: true, solution: false -->
```python
assert isprime(1) == 'not prime'
```