Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select  **Restart** and **Clear All Outputs**, then select **Run All**).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Andrea Rios"
COLLABORATORS = "BOCCHI"
```

---

# **Problem Set 4** 
### This is a problem set.  You can work on it with the help of your classmates, Google, or ask me any questions you would like. 
### I encourage discussion during these problem sets.  Todays problem set is relatively short.  


```python
import numpy as np
from matplotlib import pyplot as plt
```

Output:
```
KeyboardInterrupt
```

1. Write a for loop that prints a countdown starting from the number **n** to the number 1. For your test case, set n = 10 

<!-- grade_id: cell-c7608d833c216217, points: 3, grade: true, solution: true -->
```python
n = 10
for i in range (n,0,-1):
    print(i)
```

```python
#this question is manually graded 
assert 1 == 1
```

2. Given a color_list = ['yellow', 'green'] and a fruit_list = ['banana', 'orange', 'mango', 'lemon'],  create a new list **color_fruit** with 8 elements that combines each color with each fruit. All the yellow fruits should come first and all the green fruits should come second.  The fruits should be listed in the original order. Make it pretty by using a space between the color and fruit - ' ' 

- Hint: Loops can be nested   

<!-- grade_id: cell-e4fbcf7ef498f421, points: 0, grade: false, solution: true -->
```python
color_list = ['yellow', 'green']
fruit_list = ['banana', 'orange', 'mango', 'lemon']
color_fruit = []
for color in color_list:
    for fruit in fruit_list:
        color_fruit.append(color+' '+fruit)
        print(color_fruit)
```

<!-- grade_id: cell-e63176dfd244677b, points: 5, grade: true, solution: false -->
```python
assert len(color_fruit) == 8
assert color_fruit[0] == 'yellow banana'
assert color_fruit[6] == 'green mango'
```

3. Approximating $\pi$.  There are a number of approximations to the number $\pi$ that have been developed over the years.   

Leibniz's formula for approximating $\pi$ is

$$ \pi = 4 - 4/3 + 4/5 - 4/7 + 4/9 - 4/11 + 4/13 ....   $$

$$ \pi = \sum_{n=0}^{\infty}{(-1)^n\frac{4}{2n+1}} $$

Note two things to pay attention to about this sum. 
* The first term (n=0) is equivalent to 4/1.  The second term is 4/3. So, we are counting 1,3,5,7,9,11.... for the denominator.  That is where the 2*n+1 comes from 
* We are alternating between adding and subtracting numbers.  That is where the $(-1)^n$ comes from.    

a. Write a for loop that will compute all the approximations to $\pi$ starting using N terms.  For our test we will use N =20.   

At each iterate of the loop, calculate the approximate value of pi and then the error in the approximation 
Save the error to a list called **error** which is the difference between the exact value of $\pi$ and your approximation.  

* Hint : To get the exact value of pi to compute error use `np.pi`
  


<!-- grade_id: cell-9e1c821ed86106b8, points: 0, grade: false, solution: true -->
```python
import numpy as np
N = 20
pi_approx = 0 
error = []
for n in range (N):
    pi_approx += 4 * ((-1)**n)/(2*n+1)
    err = abs(pi_approx - np.pi)
    error.append(err)

```

<!-- grade_id: cell-d6ac644ce5e4b39e, points: 5, grade: true, solution: false -->
```python
assert np.abs(error[0]) == 4-np.pi
assert np.abs(error[4]) < 0.2
assert np.abs(error[-1]) < 0.05
assert len(error) == 20  
```

b. Make a bar plot of the error.  The x value should be the number of terms (starting at 1) and y-axis should be the value of error.  The x-axis label should be error and the x axis label should be N.  There should be an x-tick corresponding to each value of N.  There should be grid lines to make it easy to read y-values

<!-- grade_id: cell-51b21ad402823dce, points: 3, grade: true, solution: true -->
```python
# YOUR CODE HERE
plt.bar(np.arange(1,N+1), error)
```

4. The function f(x,s) is defined as 

$$ \Large f(x,s) = e^\frac{-{x^2}}{s^2} $$

Don't forget to that x and s must be squared, and there is a minus sign in front of x.  

Make a plot of f(x,s) for x ranging from -3 to 3 (inclusive) in steps of 0.1.  The plot should show the function f, for s= 1 and s = 0.5 using different color lines. 

To draw two lines with a label, make two `plt.plot` statements like this:

        plt.plot(x,f,'r-',label = 's = 1')

if you attach a label to the plot you can get a legend using, 

        plt.legend()

Label the x axis with the letter 'x' and the y axis with the letter 'f'

<!-- grade_id: cell-04e10dcac227b559, points: 4, grade: true, solution: true -->
```python

# YOUR CODE HERE
def f(x,s):
    f = np.exp(-x**2/s**2)
    return f
x=np.arange(-3,3.1,0.1)
f1=f(x,s=1)
f2 = f(x,s=0.5)
plt.plot(x,f1,label='s=1')
plt.plot(x,f2,label='s=0.5')
plt.legend()
plt.xlabel('x',fontsize=12)
plt.ylabel('y',fontsize =12)

```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[34], line 8
      6 f1=f(x,s=1)
      7 f2 = f(x,s=0.5)
----> 8 plt.plot(x,f1,label='s=1')
      9 plt.plot(x,f2,label='s=0.5')
     10 plt.legend()

NameError: name 'plt' is not defined
```

```python

```

```python

```