Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select  **Restart** and **Clear All Outputs**, then select **Run All**).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Mariela Martinez"
COLLABORATORS = "ChatGPT & Google"
```

---

# Problem Set 6 
### This is a problem set.  You can work on it with the help of your classmates, Google, or ask me any questions you would like. 
### This problem set is take home.   

```python
import numpy as np
from numpy import random 
rng = random.default_rng(seed = 121)
from sound import pop_scales, make_tone, play_sound
```

```python
!pip install sounddevice

```

Output:
```
Requirement already satisfied: sounddevice in /Users/tailai/anaconda3/envs/nbgrader/lib/python3.13/site-packages (0.5.3)
Requirement already satisfied: CFFI>=1.0 in /Users/tailai/anaconda3/envs/nbgrader/lib/python3.13/site-packages (from sounddevice) (1.17.1)
Requirement already satisfied: pycparser in /Users/tailai/anaconda3/envs/nbgrader/lib/python3.13/site-packages (from CFFI>=1.0->sounddevice) (2.23)
```

1. Write a function **play_scale** that makes and plays the **first five notes** composing one octave in either ascending or descending order. 

I have imported all the functions in sound.py so you can call them.  If you are not able to, its because sound.py is not in the same folder as this assignment.  

The function `pop_scale()` in the module `sound` returns a dictionary with all the frequencies. You can get a list of all the frequencies in the first octave using the keyword 'frequencies' 

The first note A has a frequency of 440 in the first octave. In the second octave at 880, and all the other frequencies also double. The third octave starts at double that frequency or 1760 (4*440), and all the other frequencies are multiplied by 4.  So for any note we can get the equivalent frequency at any octave using

    frequency = frequency*2**(n_octave-1)


The function **play_scale** should have four input arguments, one is the octave (**n_octave**), the second is the **duration** of each note, and the third is an optional **volume** which defaults to 0.05, and the fourth is an optional order which is a string which is either 'ascend' or 'descend', and defaults to 'ascend'. If an invalid order is given it should not play anything and should print out "Invalid Order"
 

<!-- grade_id: cell-5092c138c16fcb4b, points: 8, grade: true, solution: true -->
```python

# YOUR CODE HERE
from sound import pop_scales, make_tone, play_sound
import numpy as np

def play_scale(n_octave, duration, volume=0.05, order='ascend'):
    scale = pop_scales()
    base_freqs = scale['frequencies']
    first_five = base_freqs[:5]

    factor = 2 ** (n_octave - 1)
    octave_notes = [f * factor for f in first_five]

    if order == 'ascend':
        notes = octave_notes
    elif order == 'descend':
        notes = octave_notes[::-1]
    else:
        print("Invalid Order")
        return

    for freq in notes:
        tone = make_tone(freq, duration)
        if tone.size == 0:
            continue
        play_sound(tone, volume)
```

```python
#should play ascending order 2nd octave
play_scale(2,0.5)

```

```python
#should play descending order 1st octave
play_scale(1,0.5,order='descend')
```

```python
#should play descending order louder volume 
play_scale(2,0.5,volume = 0.15,order='descend')
```

```python
#should not play anything and should print "Invalid Order"
play_scale(2,0.5,order = 'x')
```

Output:
```
Invalid Order
```

2. Make a function **random_scales** which randomly plays an ascending or descending scale in a random octave chosen from the first, second or 3rd octave.  

- You should make use of **play_scale** that you created in problem 1.  

- **random_scale** should get as input arguments the **duration**, and an optional parameters for volume which defaults to 0.05.

- It should choose **n_octave** randomly from 1,2,3.  

- The function should play the scale and return a string with the type of scale of played, 'ascend' or 'descend'

- The function should also return which octave was played

Hint: use the random number generator `rng` to randomly select the type of scale and octave to play

<!-- grade_id: cell-7786134923d7959f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def random_scale(duration, volume=0.05):
    n_octave = rng.integers(1, 4)  # 1 ≤ n < 4 → gives 1,2,3
    order = 'ascend' if rng.integers(0, 2) == 0 else 'descend'
    play_scale(n_octave, duration, volume, order)
    return order, n_octave
```

<!-- grade_id: cell-67eb26600fce895c, points: 4, grade: true, solution: false -->
```python
rng = random.default_rng(seed = 1211)
order, n_oct = random_scale(0.5,0.2)
print(n_oct)
assert n_oct in [1,2,3]
assert n_oct == 2
assert order == 'descend'
```

Output:
```
3
AssertionError                            Traceback (most recent call last)
Cell In[10], line 5
      3 print(n_oct)
      4 assert n_oct in [1,2,3]
----> 5 assert n_oct == 2
      6 assert order == 'descend'

AssertionError:
```

3. Copy and paste your function from Problem Set 5, **get_response** that queries the subject with a prompt and checks if the response is among the correct keys. They should be prompted to respond again if they type any other key. You may use my function from the answer key.   

<!-- grade_id: cell-d6d894aebd7e7bec, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
import numpy as np

def get_response(prompt, validresponses):
    while True:
        response = input(prompt).strip()
        
        if isinstance(validresponses, np.ndarray):
            valid = np.any(validresponses == response)
        else:
            valid = response in validresponses
        
        if valid:
            print(response)
            return response
        else:
            print("Invalid Response, Try Again")
```

4. Write a function called **trial** that uses **random_scale** and **get_response** to play a random scale, then ask for a response of a for ascending and d for descending.

The subject should response with they key `a` for ascending and `d` for descending.  A good prompt might be

        'Respond a for ascending and d for descending: '

The function should evaluate if the response is correct or incorrect and create a variable **accuracy** which is 1 if they get it right and 0 if they get it wrong.

The function should return **accuracy**, **response**, and **order**

The input arguments for trial are **duration** and **volume**.  **volume** should have a default value of 0.05.   

<!-- grade_id: cell-c15eff5cb28be0f1, points: 6, grade: true, solution: true -->
```python
# YOUR CODE HERE
def trial(duration, volume=0.05):
    order, n_oct = random_scale(duration, volume)
    response = get_response("Respond a for ascending and d for descending: ", ['a', 'd'])
    accuracy = 1 if (response == 'a' and order == 'ascend') or (response == 'd' and order == 'descend') else 0
    return accuracy, response, order
```

```python
accuracy, response, order = trial(0.5)
print(accuracy)
```

Output:
```
Cell In[11], line 6, in get_response(prompt, validresponses)
      4 def get_response(prompt, validresponses):
      5     while True:
----> 6         response = input(prompt).strip()
      8         if isinstance(validresponses, np.ndarray):
      9             valid = np.any(validresponses == response)

File ~/anaconda3/envs/nbgrader/lib/python3.13/site-packages/ipykernel/kernelbase.py:1274, in Kernel.raw_input(self, prompt)
   1272 if not self._allow_stdin:
   1273     msg = "raw_input was called, but this frontend does not support input requests."
-> 1274     raise StdinNotImplementedError(msg)
   1275 return self._input_request(
   1276     str(prompt),
   1277     self._parent_ident["shell"],
   1278     self.get_parent("shell"),
   1279     password=False,
   1280 )

StdinNotImplementedError: raw_input was called, but this frontend does not support input requests.
```

```python

```