Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select  **Restart** and **Clear All Outputs**, then select **Run All**).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Elyse Jaeger"
COLLABORATORS = "SAITAMA"
```

---

# **Problem Set 2**:  Random Numbers, Array Methods, Functions

* This is a problem set.  You can work on it with the help of your classmates, Google, or ask us any questions you would like.  
* I encourage discussion during these problem sets.    
* In order for automatic grading and validation to work, it is very important that you make use of the exact variable names specified in the problem statement.
* I will always place these variable names that must be used for automatic grading in **bold**

```python
# your going to need numpy and a random number generator to do this assignment 
import numpy as np 
from numpy import random
rng = random.default_rng(seed = 2025)
```

1. Create an array **hexdice** with 20 random numbers with values ranging from 1 to 16 (inclusive). Sort the array **hexdice** into descending order, i.e., from largest to smallest.  The answer should still be in the array **hexdice**

<!-- grade_id: cell-2a0fce76f8a4fcc5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
hexdice = rng.integers(1, 17, 20)
hexdice_sorted = np.sort(hexdice)
hexdice = np.flip(hexdice_sorted)
```

<!-- grade_id: cell-c55d53d08cf6fea6, points: 3, grade: true, solution: false -->
```python
assert len(hexdice)== 20
assert np.argmax(hexdice) == 0 
```

2. Write a function **normalize** that takes a numpy array **x** and divides the array by the maximum value in the array and returns an array containing these normalized values.  You should start your function with,

       def normalize(x):

Don't forget to indent your code under the def statement, and to end your function with a return statement.     

<!-- grade_id: cell-b6350c0476590c1e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def normalize(x):
    normalize_x = x/np.max(x)
    return normalize_x
```

<!-- grade_id: cell-a2fb6ef5fc1c48a6, points: 4, grade: true, solution: false -->
```python
x = rng.uniform(0,10,20)
normx = normalize(x) 
assert np.max(normx) == 1 
```

3-5. It is challenging to come up with measures of excellence in individual defense in basketball.  One useful measure of individual defense is to compare the number of steals to fouls. Basketball players are often called for fouls when they try to steal the ball. There are only 4 players in NBA history including 2 active NBA players who steal more than they foul.    

Jimmy Butler is a remarkable defender by this measure. The following 3 lists have the years, games, steals, and fouls per game for every season of Jimmy Butler's career (not including his rookie year when he didn't play). Run the next code block so you have the variables available to answer problems 3-5. Answer the questions by writing Python code and using numpy functions, not just by figuring it out in your head.    

```python
years = np.arange(2012,2026)
games = [82,67,65,67,75,59,10,55,58,52,57,64,60,55]
steals = [1.0,1.9,1.8,1.6,1.9,2.0,2.4,1.8,1.8,2.1,1.6,1.8,1.3,1.7]
fouls = [1.2,1.6,1.7,1.9,1.5,1.3,1.8,1.7,1.4,1.4,1.5,1.3,1.1,1.4]
```

3. What was the highest steals per game in Jimmy Butler's career. Save this in the variable **best_steals**  What year did Jimmy Butler average the most steals? Save this in the variable **best_year**.  How many fouls did he commit per game in this year?  Save this in the variable **fouls_best_year**

<!-- grade_id: cell-1e855958e42b3792, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
best_steals = np.max(steals)
best_year = years[np.argmax(steals)]
fouls_best_year = fouls[np.argmax(steals)]
```

<!-- grade_id: cell-044d3ded57226e50, points: 1, grade: true, solution: false -->
```python
assert(best_steals > 2.2)
```

<!-- grade_id: cell-dcf744191451f41c, points: 2, grade: true, solution: false -->
```python
assert best_year == 2018
```

<!-- grade_id: cell-9bdfb06a665f56c8, points: 2, grade: true, solution: false -->
```python
assert fouls_best_year == 1.8
```

4.    Write a function `weightedaverage` that computes a weighted average.
*  input arguments are the array **data** and an array **weights**.
*  The function should multiply the **data** by **weights**.
*  Then it should sum that product using `np.sum`
*  Then divide by the sum of the **weights** also obtained using `np.sum`.

This is how you compute a *weighted average*. 

Start your function like this

    def weightedaverage(data,weights):

be sure to end the function with a `return` statement to return the weighted average value to global namespace.  

<!-- grade_id: cell-0d8820ed44c65f3f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def weightedaverage(data, weights):
    multiply = np.sum(data*weights)/np.sum(weights)
    return multiply
```

<!-- grade_id: cell-ccc53701ebbe194a, points: 2, grade: true, solution: false -->
```python
data = np.array([1,2,4])
weights = np.array([2,2,1])
wtavg = weightedaverage(data,weights)
assert wtavg == 2
```

<!-- grade_id: cell-324f1bd834a3b822, points: 2, grade: true, solution: false -->
```python
data = np.array([1,2,3,4,5])
weights = np.array([2,4,1,1,2])
wtavg = weightedaverage(data,weights)
assert wtavg == 2.7
```

5. We want to compute Jimmy Butler's career averages of steals and fouls.  The data provided are yearly averages.
If we using a simple mean or median to compute Jimmy Butler's career average we dont take into account that he played a *different number of games each year*. For example. in one year he got injured and only played in 10 games (see **games**)
* We want to *weight* the average of his yearly statistics by the number of games played each year. 
* The list **games** indicates the number of games played in each year.  
* Use your function `weightedaverage` to compute a weighted average of Jimmy's steals per game in a variable named **average_steals** and a weighted average of the number of fouls per game in a variable named **average_fouls**.  

<!-- grade_id: cell-a6cb876518bc27b9, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
games_array = np.array(games)
steals_array = np.array(steals)
fouls_array = np.array(fouls)
average_steals = weightedaverage(steals_array, games_array)
average_fouls = weightedaverage(fouls_array, games_array)
```

<!-- grade_id: cell-83ed95c7aa9c856a, points: 4, grade: true, solution: false -->
```python
assert np.isclose(average_steals,1.71,rtol=0.05)
assert np.isclose(average_fouls,1.47,rtol=0.05)

```