Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select  **Restart** and **Clear All Outputs**, then select **Run All**).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Jackson Chang"
COLLABORATORS = ""
```

---

# Problem Set 6 
### This is a problem set.  You can work on it with the help of your classmates, Google, or ask me any questions you would like. 
### This problem set is take home.   

```python
import numpy as np
from numpy import random 
rng = random.default_rng(seed = 121)
from sound import pop_scales, make_tone, play_sound
```

1. Write a function **play_scale** that makes and plays the **first five notes** composing one octave in either ascending or descending order. 

I have imported all the functions in sound.py so you can call them.  If you are not able to, its because sound.py is not in the same folder as this assignment.  

The function `pop_scale()` in the module `sound` returns a dictionary with all the frequencies. You can get a list of all the frequencies in the first octave using the keyword 'frequencies' 

The first note A has a frequency of 440 in the first octave. In the second octave at 880, and all the other frequencies also double. The third octave starts at double that frequency or 1760 (4*440), and all the other frequencies are multiplied by 4.  So for any note we can get the equivalent frequency at any octave using

    frequency = frequency*2**(n_octave-1)


The function **play_scale** should have four input arguments, one is the octave (**n_octave**), the second is the **duration** of each note, and the third is an optional **volume** which defaults to 0.05, and the fourth is an optional order which is a string which is either 'ascend' or 'descend', and defaults to 'ascend'. If an invalid order is given it should not play anything and should print out "Invalid Order"
 

<!-- grade_id: cell-5092c138c16fcb4b, points: 8, grade: true, solution: true -->
```python
# YOUR CODE HERE
print(pop_scales())
def play_scale(n_octave, duration, volume=0.05, order='ascend'):
    scales = pop_scales() #we set the dictionary to scales 
    freqs = scales['frequencies'] #we set the variable freq equal to the index defined in the dicitonary by 'frequencies'
    first_five_notes = freqs[:5] #here we call the first five notes of the frequency, or the first 5 indexs
    first_five_notes = first_five_notes * (2 ** (n_octave - 1)) #take the octave and subtract by 1 then multiply by 2 cause thats what an octave is, we then multiply by our frequencies to get the correct octave 
    if order == 'ascend':
        ordered = first_five_notes #if order is ascend then order is the first five notes 
    elif order == 'descend':
        ordered = first_five_notes[::-1] #if the order is descend the order is the first five notes backwards
    else:
        print("Invalid Order") #if the order is not either of those print invalid order 
        return
    for f in ordered: 
        tone = make_tone(f, duration) #for what we set the order of our first five notes to, we set the function make_tone for the duration and the notes
        play_sound(tone, volume) #then we use the play_sound function, with our arguments tone and volume to get the actual sound 
```

Output:
```
{'A': 440, 'Bb': 466, 'B': 494, 'C': 523, 'C#': 554, 'D': 587, 'Eb': 622, 'E': 659, 'F': 698, 'F#': 740, 'G': 784, 'Ab': 831, 'names': ['A', 'Bb', 'B', 'C', 'C#', 'D', 'Eb', 'E', 'F', 'F#', 'G', 'Ab'], 'frequencies': array([440, 466, 494, 523, 554, 587, 622, 659, 698, 740, 784, 831])}
```

```python
#should play ascending order 2nd octave
play_scale(2,0.5)

```

```python
#should play descending order 1st octave
play_scale(1,0.5,order='descend')
```

```python
#should play descending order louder volume 
play_scale(2,0.5,volume = 0.15,order='descend')
```

```python
#should not play anything and should print "Invalid Order"
play_scale(2,0.5,order = 'x')
```

Output:
```
Invalid Order
```

2. Make a function **random_scales** which randomly plays an ascending or descending scale in a random octave chosen from the first, second or 3rd octave.  

- You should make use of **play_scale** that you created in problem 1.  

- **random_scale** should get as input arguments the **duration**, and an optional parameters for volume which defaults to 0.05.

- It should choose **n_octave** randomly from 1,2,3.  

- The function should play the scale and return a string with the type of scale of played, 'ascend' or 'descend'

- The function should also return which octave was played

Hint: use the random number generator `rng` to randomly select the type of scale and octave to play

<!-- grade_id: cell-7786134923d7959f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def random_scale(duration, volume=0.05):
    n_octave = rng.choice([1, 2, 3]) #randomly choose the octave from one of these indexes
    order = rng.choice(['ascend', 'descend']) #randomly choose if we ascend or descend
    play_scale(n_octave, duration, volume, order) #play_scale function, random octave, set duration, set volume, random order 
    return order, n_octave
```

<!-- grade_id: cell-67eb26600fce895c, points: 4, grade: true, solution: false -->
```python
rng = random.default_rng(seed = 1211)
order, n_oct = random_scale(0.5,0.2)
print(n_oct)
assert n_oct in [1,2,3]
assert n_oct == 2
assert order == 'descend'
```

Output:
```
3
AssertionError                            Traceback (most recent call last)
Cell In[9], line 5
      3 print(n_oct)
      4 assert n_oct in [1,2,3]
----> 5 assert n_oct == 2
      6 assert order == 'descend'

AssertionError:
```

3. Copy and paste your function from Problem Set 5, **get_response** that queries the subject with a prompt and checks if the response is among the correct keys. They should be prompted to respond again if they type any other key. You may use my function from the answer key.   

<!-- grade_id: cell-d6d894aebd7e7bec, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
def get_response(prompt, validresponses):
    while True:
        response = input(prompt)
        if response in validresponses:
            print(response)
            return response
        else:
            print("Invalid Response, Try Again")
```

4. Write a function called **trial** that uses **random_scale** and **get_response** to play a random scale, then ask for a response of a for ascending and d for descending.

The subject should response with they key `a` for ascending and `d` for descending.  A good prompt might be

        'Respond a for ascending and d for descending: '

The function should evaluate if the response is correct or incorrect and create a variable **accuracy** which is 1 if they get it right and 0 if they get it wrong.

The function should return **accuracy**, **response**, and **order**

The input arguments for trial are **duration** and **volume**.  **volume** should have a default value of 0.05.   

<!-- grade_id: cell-c15eff5cb28be0f1, points: 6, grade: true, solution: true -->
```python
# YOUR CODE HERE
def trial(duration, volume=0.05):
    order, n_octave = random_scale(duration, volume) #taking the random container from the previous part 
    prompt = 'Respond a for ascending and d for descending for what you heard in the sound: ' #our response input instructions thingy 
    response = get_response(prompt, ['a', 'd']) #we use our previous get_responses function and use the valid inputs as a or d 
    if (response == 'a' and order == 'ascend') or (response == 'd' and order == 'descend'): 
        accuracy = 1 #when we type our response and we get it right the accuracy is 1 
    else:
        accuracy = 0 #if we type the wrong responses our accuracy is 0 
    return accuracy, response, order

```

```python
accuracy, response, order = trial(0.5)
print(accuracy)
```

Output:
```
Cell In[10], line 4, in get_response(prompt, validresponses)
      2 def get_response(prompt, validresponses):
      3     while True:
----> 4         response = input(prompt)
      5         if response in validresponses:
      6             print(response)

File ~/anaconda3/envs/nbgrader/lib/python3.13/site-packages/ipykernel/kernelbase.py:1274, in Kernel.raw_input(self, prompt)
   1272 if not self._allow_stdin:
   1273     msg = "raw_input was called, but this frontend does not support input requests."
-> 1274     raise StdinNotImplementedError(msg)
   1275 return self._input_request(
   1276     str(prompt),
   1277     self._parent_ident["shell"],
   1278     self.get_parent("shell"),
   1279     password=False,
   1280 )

StdinNotImplementedError: raw_input was called, but this frontend does not support input requests.
```

```python

```