Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select  **Restart** and **Clear All Outputs**, then select **Run All**).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Maia"
COLLABORATORS = "Chat gpt"
```

```python
import sys
sys.path.append('/Users/maianiemen/Desktop/Python')
import sound
```

---

# Problem Set 6 
### This is a problem set.  You can work on it with the help of your classmates, Google, or ask me any questions you would like. 
### This problem set is take home.   

```python
import numpy as np
from numpy import random 
rng = random.default_rng(seed = 121)
from sound import pop_scales, make_tone, play_sound
```

```python
scales = pop_scales()
scales['frequencies']

```

Output:
```
array([440, 466, 494, 523, 554, 587, 622, 659, 698, 740, 784, 831])
```

```python
import sounddevice as sd
sd.default.device = 1

```

```python
import numpy as np
sd.play(0.25*np.ones(30000), 44100)

```

Output:
```
File ~/anaconda3/envs/nbgrader/lib/python3.13/site-packages/sounddevice.py:909, in _StreamBase.__init__(self, kind, samplerate, blocksize, device, channels, dtype, latency, extra_settings, callback, finished_callback, clip_off, dither_off, never_drop_input, prime_output_buffers_using_stream_callback, userdata, wrap_callback)
    907     userdata = _ffi.NULL
    908 self._ptr = _ffi.new('PaStream**')
--> 909 _check(_lib.Pa_OpenStream(self._ptr, iparameters, oparameters,
    910                           samplerate, blocksize, stream_flags,
    911                           callback_ptr, userdata),
    912        f'Error opening {self.__class__.__name__}')
    914 # dereference PaStream** --> PaStream*
    915 self._ptr = self._ptr[0]

File ~/anaconda3/envs/nbgrader/lib/python3.13/site-packages/sounddevice.py:2823, in _check(err, msg)
   2820     hosterror_info = host_api, info.errorCode, hosterror_text
   2821     raise PortAudioError(errormsg, err, hosterror_info)
-> 2823 raise PortAud
```

1. Write a function **play_scale** that makes and plays the **first five notes** composing one octave in either ascending or descending order. 

I have imported all the functions in sound.py so you can call them.  If you are not able to, its because sound.py is not in the same folder as this assignment.  

The function `pop_scale()` in the module `sound` returns a dictionary with all the frequencies. You can get a list of all the frequencies in the first octave using the keyword 'frequencies' 

The first note A has a frequency of 440 in the first octave. In the second octave at 880, and all the other frequencies also double. The third octave starts at double that frequency or 1760 (4*440), and all the other frequencies are multiplied by 4.  So for any note we can get the equivalent frequency at any octave using

    frequency = frequency*2**(n_octave-1)


The function **play_scale** should have four input arguments, one is the octave (**n_octave**), the second is the **duration** of each note, and the third is an optional **volume** which defaults to 0.05, and the fourth is an optional order which is a string which is either 'ascend' or 'descend', and defaults to 'ascend'. If an invalid order is given it should not play anything and should print out "Invalid Order"
 

<!-- grade_id: cell-5092c138c16fcb4b, points: 8, grade: true, solution: true -->
```python
def play_scale(n_octave, duration, volume=0.05, order='ascend'):
    scales = pop_scales()
    base_freqs = scales['frequencies'][:5]
    adjusted = [f * 2**(n_octave - 1) for f in base_freqs]

    if order == 'ascend':
        final_list = adjusted
    elif order == 'descend':
        final_list = adjusted[::-1]
    else:
        print("Invalid Order")
        return

    print("Final freqs:", final_list)

    for freq in final_list:
        print("Playing freq:", freq)
        tone = make_tone(freq, duration)
        print("Tone max:", np.max(np.abs(tone)))
        play_sound(tone, volume=volume)

print("DEBUG VERSION LOADED")


```

Output:
```
DEBUG VERSION LOADED
```

```python
#should play ascending order 2nd octave
play_scale(2,0.5)

```

Output:
```
Final freqs: [np.int64(880), np.int64(932), np.int64(988), np.int64(1046), np.int64(1108)]
Playing freq: 880
Tone max: 0.9999999974623486
File ~/anaconda3/envs/nbgrader/lib/python3.13/site-packages/sounddevice.py:909, in _StreamBase.__init__(self, kind, samplerate, blocksize, device, channels, dtype, latency, extra_settings, callback, finished_callback, clip_off, dither_off, never_drop_input, prime_output_buffers_using_stream_callback, userdata, wrap_callback)
    907     userdata = _ffi.NULL
    908 self._ptr = _ffi.new('PaStream**')
--> 909 _check(_lib.Pa_OpenStream(self._ptr, iparameters, oparameters,
    910                           samplerate, blocksize, stream_flags,
    911                           callback_ptr, userdata),
    912        f'Error opening {self.__class__.__name__}')
    914 # dereference PaStream** --> PaStream*
    915 self._ptr = self._ptr[0]

File ~/anaconda3/envs/nbgrader/lib/python3.13/site-packages/sounddevice.py:2823, in _check(err, msg)
   2820     hoste
```

```python
#should play descending order 1st octave
play_scale(1,0.5,order='descend')
```

Output:
```
Final freqs: [np.int64(554), np.int64(523), np.int64(494), np.int64(466), np.int64(440)]
Playing freq: 554
Tone max: 0.9999999974623486
File ~/anaconda3/envs/nbgrader/lib/python3.13/site-packages/sounddevice.py:909, in _StreamBase.__init__(self, kind, samplerate, blocksize, device, channels, dtype, latency, extra_settings, callback, finished_callback, clip_off, dither_off, never_drop_input, prime_output_buffers_using_stream_callback, userdata, wrap_callback)
    907     userdata = _ffi.NULL
    908 self._ptr = _ffi.new('PaStream**')
--> 909 _check(_lib.Pa_OpenStream(self._ptr, iparameters, oparameters,
    910                           samplerate, blocksize, stream_flags,
    911                           callback_ptr, userdata),
    912        f'Error opening {self.__class__.__name__}')
    914 # dereference PaStream** --> PaStream*
    915 self._ptr = self._ptr[0]

File ~/anaconda3/envs/nbgrader/lib/python3.13/site-packages/sounddevice.py:2823, in _check(err, msg)
   2820     hosterr
```

```python
#should play descending order louder volume 
play_scale(2,0.5,volume = 0.15,order='descend')
```

Output:
```
Final freqs: [np.int64(1108), np.int64(1046), np.int64(988), np.int64(932), np.int64(880)]
Playing freq: 1108
Tone max: 0.9999999974623486
File ~/anaconda3/envs/nbgrader/lib/python3.13/site-packages/sounddevice.py:909, in _StreamBase.__init__(self, kind, samplerate, blocksize, device, channels, dtype, latency, extra_settings, callback, finished_callback, clip_off, dither_off, never_drop_input, prime_output_buffers_using_stream_callback, userdata, wrap_callback)
    907     userdata = _ffi.NULL
    908 self._ptr = _ffi.new('PaStream**')
--> 909 _check(_lib.Pa_OpenStream(self._ptr, iparameters, oparameters,
    910                           samplerate, blocksize, stream_flags,
    911                           callback_ptr, userdata),
    912        f'Error opening {self.__class__.__name__}')
    914 # dereference PaStream** --> PaStream*
    915 self._ptr = self._ptr[0]

File ~/anaconda3/envs/nbgrader/lib/python3.13/site-packages/sounddevice.py:2823, in _check(err, msg)
   2820     host
```

```python
#should not play anything and should print "Invalid Order"
play_scale(2,0.5,order = 'x')
```

Output:
```
Invalid Order
```

```python
import sounddevice as sd
sd.query_devices()

```

Output:
```
0 Tailais iPhone Microphone, Core Audio (1 in, 0 out)
* 1 MacBook Air Microphone, Core Audio (1 in, 0 out)
  2 MacBook Air Speakers, Core Audio (0 in, 2 out)
  3 Microsoft Teams Audio, Core Audio (2 in, 2 out)
  4 ZoomAudioDevice, Core Audio (2 in, 2 out)
```

2. Make a function **random_scales** which randomly plays an ascending or descending scale in a random octave chosen from the first, second or 3rd octave.  

- You should make use of **play_scale** that you created in problem 1.  

- **random_scale** should get as input arguments the **duration**, and an optional parameters for volume which defaults to 0.05.

- It should choose **n_octave** randomly from 1,2,3.  

- The function should play the scale and return a string with the type of scale of played, 'ascend' or 'descend'

- The function should also return which octave was played

Hint: use the random number generator `rng` to randomly select the type of scale and octave to play

<!-- grade_id: cell-7786134923d7959f, points: 0, grade: false, solution: true -->
```python
def random_scale(duration, volume=0.05):
    n_oct = rng.integers(1, 4)
    order = rng.choice(['ascend', 'descend'])
    play_scale(n_oct, duration, volume=volume, order=order)
    return order, n_oct

```

<!-- grade_id: cell-67eb26600fce895c, points: 4, grade: true, solution: false -->
```python
rng = random.default_rng(seed = 1211)
order, n_oct = random_scale(0.5,0.2)
print(n_oct)
assert n_oct in [1,2,3]
assert n_oct == 2
assert order == 'descend'
```

Output:
```
Final freqs: [np.int64(1760), np.int64(1864), np.int64(1976), np.int64(2092), np.int64(2216)]
Playing freq: 1760
Tone max: 0.9999999974623486
File ~/anaconda3/envs/nbgrader/lib/python3.13/site-packages/sounddevice.py:909, in _StreamBase.__init__(self, kind, samplerate, blocksize, device, channels, dtype, latency, extra_settings, callback, finished_callback, clip_off, dither_off, never_drop_input, prime_output_buffers_using_stream_callback, userdata, wrap_callback)
    907     userdata = _ffi.NULL
    908 self._ptr = _ffi.new('PaStream**')
--> 909 _check(_lib.Pa_OpenStream(self._ptr, iparameters, oparameters,
    910                           samplerate, blocksize, stream_flags,
    911                           callback_ptr, userdata),
    912        f'Error opening {self.__class__.__name__}')
    914 # dereference PaStream** --> PaStream*
    915 self._ptr = self._ptr[0]

File ~/anaconda3/envs/nbgrader/lib/python3.13/site-packages/sounddevice.py:2823, in _check(err, msg)
   2820     h
```

3. Copy and paste your function from Problem Set 5, **get_response** that queries the subject with a prompt and checks if the response is among the correct keys. They should be prompted to respond again if they type any other key. You may use my function from the answer key.   

<!-- grade_id: cell-d6d894aebd7e7bec, points: 2, grade: true, solution: true -->
```python
def get_response(prompt,validresponses):
    validresponses = np.array(validresponses)
    invalid = True
    while invalid:
        response = input(prompt)
        if np.any(validresponses == response):

            print(response)
            invalid = False
        else:
            print('Invalid Response, Try Again')
    return response 
```

4. Write a function called **trial** that uses **random_scale** and **get_response** to play a random scale, then ask for a response of a for ascending and d for descending.

The subject should response with they key `a` for ascending and `d` for descending.  A good prompt might be

        'Respond a for ascending and d for descending: '

The function should evaluate if the response is correct or incorrect and create a variable **accuracy** which is 1 if they get it right and 0 if they get it wrong.

The function should return **accuracy**, **response**, and **order**

The input arguments for trial are **duration** and **volume**.  **volume** should have a default value of 0.05.   

<!-- grade_id: cell-c15eff5cb28be0f1, points: 6, grade: true, solution: true -->
```python
def trial(duration, volume=0.05):
    order, n_oct = random_scale(duration, volume)
    response = get_response("Respond a for ascending and d for descending: ", ['a','d'])
    
    user_order = 'ascend' if response == 'a' else 'descend'
    accuracy = 1 if user_order == order else 0
    
    return accuracy, response, order

```

```python
accuracy, response, order = trial(0.5)
print(accuracy)
```

Output:
```
Final freqs: [np.int64(1108), np.int64(1046), np.int64(988), np.int64(932), np.int64(880)]
Playing freq: 1108
Tone max: 0.9999999974623486
File ~/anaconda3/envs/nbgrader/lib/python3.13/site-packages/sounddevice.py:909, in _StreamBase.__init__(self, kind, samplerate, blocksize, device, channels, dtype, latency, extra_settings, callback, finished_callback, clip_off, dither_off, never_drop_input, prime_output_buffers_using_stream_callback, userdata, wrap_callback)
    907     userdata = _ffi.NULL
    908 self._ptr = _ffi.new('PaStream**')
--> 909 _check(_lib.Pa_OpenStream(self._ptr, iparameters, oparameters,
    910                           samplerate, blocksize, stream_flags,
    911                           callback_ptr, userdata),
    912        f'Error opening {self.__class__.__name__}')
    914 # dereference PaStream** --> PaStream*
    915 self._ptr = self._ptr[0]

File ~/anaconda3/envs/nbgrader/lib/python3.13/site-packages/sounddevice.py:2823, in _check(err, msg)
   2820     host
```

```python

```