Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select  **Restart** and **Clear All Outputs**, then select **Run All**).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Yuqiang Wang"
COLLABORATORS = ""
```

---

# Problem Set 6 
### This is a problem set.  You can work on it with the help of your classmates, Google, or ask me any questions you would like. 
### This problem set is take home.   

```python
import numpy as np
from numpy import random 
rng = random.default_rng(seed = 121)
from sound import pop_scales, make_tone, play_sound
```

1. Write a function **play_scale** that makes and plays the **first five notes** composing one octave in either ascending or descending order. 

I have imported all the functions in sound.py so you can call them.  If you are not able to, its because sound.py is not in the same folder as this assignment.  

The function `pop_scale()` in the module `sound` returns a dictionary with all the frequencies. You can get a list of all the frequencies in the first octave using the keyword 'frequencies' 

The first note A has a frequency of 440 in the first octave. In the second octave at 880, and all the other frequencies also double. The third octave starts at double that frequency or 1760 (4*440), and all the other frequencies are multiplied by 4.  So for any note we can get the equivalent frequency at any octave using

    frequency = frequency*2**(n_octave-1)


The function **play_scale** should have four input arguments, one is the octave (**n_octave**), the second is the **duration** of each note, and the third is an optional **volume** which defaults to 0.05, and the fourth is an optional order which is a string which is either 'ascend' or 'descend', and defaults to 'ascend'. If an invalid order is given it should not play anything and should print out "Invalid Order"
 

<!-- grade_id: cell-5092c138c16fcb4b, points: 8, grade: true, solution: true -->
```python
def play_scale(n_octave, duration, volume=0.05, order='ascend'):
    scales = pop_scales()
    freqs = scales['frequencies'] 
    freqs = freqs[0:5]
    freqs = [f * 2**(n_octave - 1) for f in freqs]
    if order == 'ascend':
        for f in freqs:
            tone = make_tone(f, duration)
            play_sound(tone, volume=volume)
    elif order == 'descend':
        for f in reversed(freqs):
            tone = make_tone(f, duration)
            play_sound(tone, volume=volume)
    else:
        print("Invalid Order")

```

```python
#should play ascending order 2nd octave
play_scale(2,0.5)

```

Output:
```
File ~/.local/lib/python3.10/site-packages/sounddevice.py:2735, in _get_stream_parameters(kind, device, channels, dtype, latency, extra_settings, samplerate)
   2732 if samplerate is None:
   2733     samplerate = default.samplerate
-> 2735 info = query_devices(device)
   2736 if channels is None:
   2737     channels = info['max_' + kind + '_channels']

File ~/.local/lib/python3.10/site-packages/sounddevice.py:572, in query_devices(device, kind)
    570 info = _lib.Pa_GetDeviceInfo(device)
    571 if not info:
--> 572     raise PortAudioError(f'Error querying device {device}')
    573 assert info.structVersion == 2
    574 name_bytes = _ffi.string(info.name)

PortAudioError: Error querying device -1
```

```python
#should play descending order 1st octave
play_scale(1,0.5,order='descend')
```

Output:
```
File ~/.local/lib/python3.10/site-packages/sounddevice.py:2735, in _get_stream_parameters(kind, device, channels, dtype, latency, extra_settings, samplerate)
   2732 if samplerate is None:
   2733     samplerate = default.samplerate
-> 2735 info = query_devices(device)
   2736 if channels is None:
   2737     channels = info['max_' + kind + '_channels']

File ~/.local/lib/python3.10/site-packages/sounddevice.py:572, in query_devices(device, kind)
    570 info = _lib.Pa_GetDeviceInfo(device)
    571 if not info:
--> 572     raise PortAudioError(f'Error querying device {device}')
    573 assert info.structVersion == 2
    574 name_bytes = _ffi.string(info.name)

PortAudioError: Error querying device -1
```

```python
#should play descending order louder volume 
play_scale(2,0.5,volume = 0.15,order='descend')
```

Output:
```
File ~/.local/lib/python3.10/site-packages/sounddevice.py:2735, in _get_stream_parameters(kind, device, channels, dtype, latency, extra_settings, samplerate)
   2732 if samplerate is None:
   2733     samplerate = default.samplerate
-> 2735 info = query_devices(device)
   2736 if channels is None:
   2737     channels = info['max_' + kind + '_channels']

File ~/.local/lib/python3.10/site-packages/sounddevice.py:572, in query_devices(device, kind)
    570 info = _lib.Pa_GetDeviceInfo(device)
    571 if not info:
--> 572     raise PortAudioError(f'Error querying device {device}')
    573 assert info.structVersion == 2
    574 name_bytes = _ffi.string(info.name)

PortAudioError: Error querying device -1
```

```python
#should not play anything and should print "Invalid Order"
play_scale(2,0.5,order = 'x')
```

Output:
```
Invalid Order
```

2. Make a function **random_scales** which randomly plays an ascending or descending scale in a random octave chosen from the first, second or 3rd octave.  

- You should make use of **play_scale** that you created in problem 1.  

- **random_scale** should get as input arguments the **duration**, and an optional parameters for volume which defaults to 0.05.

- It should choose **n_octave** randomly from 1,2,3.  

- The function should play the scale and return a string with the type of scale of played, 'ascend' or 'descend'

- The function should also return which octave was played

Hint: use the random number generator `rng` to randomly select the type of scale and octave to play

<!-- grade_id: cell-7786134923d7959f, points: 0, grade: false, solution: true -->
```python
def random_scale(duration, volume=0.05):
    n_oct = rng.integers(1, 4)
    
    order = rng.choice(['ascend', 'descend'])
    
    play_scale(n_oct, duration, volume, order)
    
    return order, n_oct

```

<!-- grade_id: cell-67eb26600fce895c, points: 4, grade: true, solution: false -->
```python
rng = random.default_rng(seed = 1211)
order, n_oct = random_scale(0.5,0.2)
assert n_oct in [1,2,3]
assert order in ['ascend','descend']
```

Output:
```
File ~/.local/lib/python3.10/site-packages/sounddevice.py:2735, in _get_stream_parameters(kind, device, channels, dtype, latency, extra_settings, samplerate)
   2732 if samplerate is None:
   2733     samplerate = default.samplerate
-> 2735 info = query_devices(device)
   2736 if channels is None:
   2737     channels = info['max_' + kind + '_channels']

File ~/.local/lib/python3.10/site-packages/sounddevice.py:572, in query_devices(device, kind)
    570 info = _lib.Pa_GetDeviceInfo(device)
    571 if not info:
--> 572     raise PortAudioError(f'Error querying device {device}')
    573 assert info.structVersion == 2
    574 name_bytes = _ffi.string(info.name)

PortAudioError: Error querying device -1
```

3. Copy and paste your function from Problem Set 5, **get_response** that queries the subject with a prompt and checks if the response is among the correct keys. They should be prompted to respond again if they type any other key. You may use my function from the answer key.   

<!-- grade_id: cell-d6d894aebd7e7bec, points: 2, grade: true, solution: true -->
```python
def get_response(prompt, valid_keys):
    response = input(prompt)
    
    while response not in valid_keys:
        print("Invalid response. Please try again.")
        response = input(prompt)
    
    return response
```

4. Write a function called **trial** that uses **random_scale** and **get_response** to play a random scale, then ask for a response of a for ascending and d for descending.

The subject should response with they key `a` for ascending and `d` for descending.  A good prompt might be

        'Respond a for ascending and d for descending: '

The function should evaluate if the response is correct or incorrect and create a variable **accuracy** which is 1 if they get it right and 0 if they get it wrong.

The function should return **accuracy**, **response**, and **order**

The input arguments for trial are **duration** and **volume**.  **volume** should have a default value of 0.05.   

<!-- grade_id: cell-c15eff5cb28be0f1, points: 6, grade: true, solution: true -->
```python
def trial(duration, volume=0.05):
    order, n_oct = random_scale(duration, volume)
    
    prompt = "Respond a for ascending and d for descending: "
    response = get_response(prompt, ['a', 'd'])
    
    # ascend  a
    # descend  d
    if (order == 'ascend' and response == 'a') or \
       (order == 'descend' and response == 'd'):
        accuracy = 1
    else:
        accuracy = 0
    
    return accuracy, response, order
```

```python
accuracy, response, order = trial(0.5)
print(accuracy)
```

Output:
```
File ~/.local/lib/python3.10/site-packages/sounddevice.py:2735, in _get_stream_parameters(kind, device, channels, dtype, latency, extra_settings, samplerate)
   2732 if samplerate is None:
   2733     samplerate = default.samplerate
-> 2735 info = query_devices(device)
   2736 if channels is None:
   2737     channels = info['max_' + kind + '_channels']

File ~/.local/lib/python3.10/site-packages/sounddevice.py:572, in query_devices(device, kind)
    570 info = _lib.Pa_GetDeviceInfo(device)
    571 if not info:
--> 572     raise PortAudioError(f'Error querying device {device}')
    573 assert info.structVersion == 2
    574 name_bytes = _ffi.string(info.name)

PortAudioError: Error querying device -1
```

```python

```