Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select  **Restart** and **Clear All Outputs**, then select **Run All**).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Alexis Klein"
COLLABORATORS = "Polina, Chat"
```

---

# Problem Set 6 
### This is a problem set.  You can work on it with the help of your classmates, Google, or ask me any questions you would like. 
### This problem set is take home.   

```python
import numpy as np
from numpy import random 
rng = random.default_rng(seed = 121)
from sound import pop_scales, make_tone, play_sound
```

1. Write a function **play_scale** that makes and plays the **first five notes** composing one octave in either ascending or descending order. 

I have imported all the functions in sound.py so you can call them.  If you are not able to, its because sound.py is not in the same folder as this assignment.  

The function `pop_scale()` in the module `sound` returns a dictionary with all the frequencies. You can get a list of all the frequencies in the first octave using the keyword 'frequencies' 

The first note A has a frequency of 440 in the first octave. In the second octave at 880, and all the other frequencies also double. The third octave starts at double that frequency or 1760 (4*440), and all the other frequencies are multiplied by 4.  So for any note we can get the equivalent frequency at any octave using

    frequency = frequency*2**(n_octave-1)


The function **play_scale** should have four input arguments, one is the octave (**n_octave**), the second is the **duration** of each note, and the third is an optional **volume** which defaults to 0.05, and the fourth is an optional order which is a string which is either 'ascend' or 'descend', and defaults to 'ascend'. If an invalid order is given it should not play anything and should print out "Invalid Order"
 

<!-- grade_id: cell-5092c138c16fcb4b, points: 8, grade: true, solution: true -->
```python
# YOUR CODE HERE stop forgetting to define things
def play_scale(n_octave, duration, volume=0.05, order='ascend'):
    notes = pop_scales() # dictionary 1st octive (all the ntoes!)
    first_oct = notes['frequencies'][0:5] #1st 5 notes from 1st octave

    frequency = 2**(n_octave - 1) * first_oct # first octave from popscales() * 2^(whatever given octave)

    if order == 'ascend': #default
        final_freqs = frequency
    elif order == 'descend':
        final_freqs = frequency[::-1] #reverse order (like on exam2)
    else:
        print("Invalid order")
        return # have to put this here to stop the function if there's an invalid order... Can't use 'break' b/c not a loop

    for f in final_freqs:
        tone = make_tone(f, duration) # maketone() returns pure tone (sine wave of whatever frequency is being played)
        play_sound(tone, volume=volume) # play_sound() applies volume of whatever current tone at given volume.
```

```python
#should play ascending order 2nd octave
play_scale(2,0.5)

```

```python
#should play descending order 1st octave
play_scale(1,0.5,order='descend')
```

```python
#should play descending order louder volume 
play_scale(2,0.5,volume = 0.15,order='descend')
```

```python
#should not play anything and should print "Invalid Order"
play_scale(2,0.5,order = 'x')
```

Output:
```
Invalid order
```

2. Make a function **random_scales** which randomly plays an ascending or descending scale in a random octave chosen from the first, second or 3rd octave.  

- You should make use of **play_scale** that you created in problem 1.  

- **random_scale** should get as input arguments the **duration**, and an optional parameters for volume which defaults to 0.05.

- It should choose **n_octave** randomly from 1,2,3.  

- The function should play the scale and return a string with the type of scale of played, 'ascend' or 'descend'

- The function should also return which octave was played

Hint: use the random number generator `rng` to randomly select the type of scale and octave to play

<!-- grade_id: cell-7786134923d7959f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def random_scale(duration, volume=0.05):
    n_octave = rng.integers(1, 4) #random values between 1 & 3 (as requested)

    random_order = rng.integers(0,2) #will pick either 0 or 1, we can assign ascend/descend to them
    if random_order == 0:
        order = 'ascend' #this is also the default
    else:
        order = 'descend'
    
    play_scale(n_octave, duration, volume = volume, order = order) #MAKE SURE TO DEFINE EVERYTHING OTHERWISE YOU'LL HAVE ISSUES LATER
    print(order, n_octave)
    return order, n_octave

#This also works.
#def random_scale(duration, volume=0.05):
    #octaves = np.array([1,2,3]) #convert into numpy array so u can use np.permutation
    #n_octave = rng.permutation(octaves)[0] # rng.permutation() shuffles the numbers, [0] returns the first value.
    
    #orders = ['ascend', 'descend']
    #order = rng.permutation(orders)[0] #again, shuffles 'asc' & 'desc', then picks whatever the first value is
    
    #play_scale(n_octave, duration, volume, order) #calling function play_scale()
    #print(order, n_octave)
    #return order, n_octave

```

<!-- grade_id: cell-67eb26600fce895c, points: 4, grade: true, solution: false -->
```python
rng = random.default_rng(seed = 1211)
order, n_oct = random_scale(0.5,0.2)
print(n_oct)
assert n_oct in [1,2,3]
assert n_oct == 2
assert order == 'descend'
```

Output:
```
ascend 3
3
AssertionError                            Traceback (most recent call last)
Cell In[9], line 5
      3 print(n_oct)
      4 assert n_oct in [1,2,3]
----> 5 assert n_oct == 2
      6 assert order == 'descend'

AssertionError:
```

3. Copy and paste your function from Problem Set 5, **get_response** that queries the subject with a prompt and checks if the response is among the correct keys. They should be prompted to respond again if they type any other key. You may use my function from the answer key.   

<!-- grade_id: cell-d6d894aebd7e7bec, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
def get_response(prompt, validresponses):
    response = input(prompt)
    is_valid = False
    while is_valid ==False:
        for valid in validresponses:
            if response ==valid:
                is_valid = True
                break
        if is_valid ==False:
            print("Invalid response, try again")
            response = input(prompt)
    print(response)
    return response
```

4. Write a function called **trial** that uses **random_scale** and **get_response** to play a random scale, then ask for a response of a for ascending and d for descending.

The subject should response with they key `a` for ascending and `d` for descending.  A good prompt might be

        'Respond a for ascending and d for descending: '

The function should evaluate if the response is correct or incorrect and create a variable **accuracy** which is 1 if they get it right and 0 if they get it wrong.

The function should return **accuracy**, **response**, and **order**

The input arguments for trial are **duration** and **volume**.  **volume** should have a default value of 0.05.   

<!-- grade_id: cell-c15eff5cb28be0f1, points: 6, grade: true, solution: true -->
```python
# YOUR CODE HERE
def trial(duration, volume=0.05):
    order, n_octave = random_scale(duration, volume) # random scale!!!!!!

    prompt = "Respond 'a' for ascending and 'd' for descending: "#defining prompt for get_response
    validresponses = ['a', 'd'] #defining valid responses for get_response
    response = get_response(prompt, validresponses) #call get response to get yer responses

    if order == 'ascend' and response == 'a': #if the order chosen by random_scale is ascend & participant inputs 'a', get accuracy point
        accuracy = 1
    elif order == 'descend' and response == 'd': #same goes for descend
        accuracy = 1
    else:
        accuracy = 0 #otherwise they get 0

    print('accuracy:', accuracy, ' response:', response, ' order:', order) #print stuff
    return accuracy, response, order #return statement

```

```python
accuracy, response, order = trial(0.5)
print(accuracy)
```

Output:
```
descend 2
Cell In[10], line 3, in get_response(prompt, validresponses)
      2 def get_response(prompt, validresponses):
----> 3     response = input(prompt)
      4     is_valid = False
      5     while is_valid ==False:

File ~/anaconda3/envs/nbgrader/lib/python3.13/site-packages/ipykernel/kernelbase.py:1274, in Kernel.raw_input(self, prompt)
   1272 if not self._allow_stdin:
   1273     msg = "raw_input was called, but this frontend does not support input requests."
-> 1274     raise StdinNotImplementedError(msg)
   1275 return self._input_request(
   1276     str(prompt),
   1277     self._parent_ident["shell"],
   1278     self.get_parent("shell"),
   1279     password=False,
   1280 )

StdinNotImplementedError: raw_input was called, but this frontend does not support input requests.
```