Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select  **Restart** and **Clear All Outputs**, then select **Run All**).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# **Problem Set 1**:  List, Arrays, and Indexing Methods

### This is a problem set.  You can work on it with the help of your classmates, Google, or ask me any questions you would like. 
### There are some new methods introduced in the problem set. 
### I encourage discussion during these problem sets.   
### The problem set is due at 2:50 pm.   
### However, you may continue to work on problems you dont finish, or you want to correct up to Wednesday. You will get 1/2 credit for that additional work.   
### VERY IMPORTANT
### In order for automatic grading and validation to work, it is very important that you make use of the exact variable names specified in the problem statement.
### I will always place these variable names that must be used for automatic grading in **bold**

```python
# your going to need numpy to do this assignment 
import numpy as np 
#frankly, in this class, you will always need numpy.  
```

1. Create a list called **food** with the following elements 'onion', 'egg', 'bacon', 'tomato', 'ham', 'cheese'. 

<!-- grade_id: cell-2a0fce76f8a4fcc5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-c55d53d08cf6fea6, points: 3, grade: true, solution: false -->
```python
assert len(food)== 6
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[4], line 1
----> 1 assert len(food)== 6

NameError: name 'food' is not defined
```

2. Use *indexing* into the list **food** to make a new list **meat** containing only 'bacon' and 'ham'. Again, using indexing, make a list **veggie** containing 'tomato' and 'onion' in that order. 

<!-- grade_id: cell-b6350c0476590c1e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-a2fb6ef5fc1c48a6, points: 2, grade: true, solution: false -->
```python
assert meat[0] == 'bacon'
assert len(meat) == 2
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[6], line 1
----> 1 assert meat[0] == 'bacon'
      2 assert len(meat) == 2

NameError: name 'meat' is not defined
```

<!-- grade_id: cell-6a99beb4a531ef02, points: 2, grade: true, solution: false -->
```python
assert veggie[0] == 'tomato'
assert len(veggie) == 2
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[7], line 1
----> 1 assert veggie[0] == 'tomato'
      2 assert len(veggie) == 2

NameError: name 'veggie' is not defined
```

3. Create a numpy array called **numbers** containing the numbers 6,24,13,16,8,10,21 in that order. Create an array called **skip1** that starts at element with index 0 and takes every other element of the array **numbers**, by using the `start:stop:step` syntax to index into the array.

<!-- grade_id: cell-c9c71eda64dac60b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-044d3ded57226e50, points: 1, grade: true, solution: false -->
```python
assert len(numbers) == 7
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[9], line 1
----> 1 assert len(numbers) == 7

NameError: name 'numbers' is not defined
```

<!-- grade_id: cell-3bb4e599e35e1645, points: 2, grade: true, solution: false -->
```python
assert len(skip1) == 4
assert skip1[-1] == 21
assert skip1[1] == 13
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[10], line 1
----> 1 assert len(skip1) == 4
      2 assert skip1[-1] == 21
      3 assert skip1[1] == 13

NameError: name 'skip1' is not defined
```

4. Make a numpy array called **z** using the numpy function `linspace`, containing values from -100 to 100, including -100 and 100, in steps of 10.  

<!-- grade_id: cell-f95007e3ee1a7d98, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-536cfe4c84bf740a, points: 3, grade: true, solution: false -->
```python
assert np.all(np.diff(z) == 10)
### BEGIN HIDDEN TESTS
assert np.size(z) == 21
### END HIDDEN TESTS
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[12], line 1
----> 1 assert np.all(np.diff(z) == 10)
      2 ### BEGIN HIDDEN TESTS
      3 assert np.size(z) == 21

NameError: name 'z' is not defined
```

5. Create a numpy array **x** with all the odd numbers less than 20 using the numpy function `arange`.  

<!-- grade_id: cell-6bfea95ba011db9d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-925a224d993057b4, points: 3, grade: true, solution: false -->
```python
assert ~np.any(x%2 == 0)
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[14], line 1
----> 1 assert ~np.any(x%2 == 0)

NameError: name 'x' is not defined
```

6. Create an array **power** than has all of the powers of 2 from $2^1$ to $2^8$.  There are a number of ways you could tackle this problem and I will accept any of them as long as you produce the array by calculating the powers using python code. 

<!-- grade_id: cell-a6cb876518bc27b9, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-83ed95c7aa9c856a, points: 4, grade: true, solution: false -->
```python
assert len(power) == 8
assert np,max(power) == 2**8
assert np.min(power) == 2
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[16], line 1
----> 1 assert len(power) == 8
      2 assert np,max(power) == 2**8
      3 assert np.min(power) == 2

NameError: name 'power' is not defined
```