Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select  **Restart** and **Clear All Outputs**, then select **Run All**).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Yuqiang Wang"
COLLABORATORS = "Fighter D"
```

---

## **Problem Set 5** Execution Control 2

This is a problem set. You can work on it with the help of your classmates, Google, or ask me any questions you would like.

I encourage discussion during these problem sets.  

```python
import numpy as np
```

1-2.The factorial of an integer is the product of the all the integers starting from the integer going down to one.  

So, for example, in math we usually write 

$ n!  = n \times (n-1) \times (n-2)....2 \times 1 $

For example, 5! = 5 * 4 * 3 * 2 * 1 = 120 
Of course in any real program, we would use `np.prod` or `math.factorial` to compute the factorial of an integer n. 

But, we are here to suffer and learn. So you cannot used any in-built factorial function like `math.factorial`.  You are also not allowed to use `np.prod`

You can solve this problem using a `for` loop or using a `while` loop.  In this exercise you will do it both ways.  

1.  Write a function called **forfact** that take an input argument **n** and returns the factorial of n.  I want you to solve it using a **for** loop. 

<!-- grade_id: cell-7d70435eb3db2a69, points: 0, grade: false, solution: true -->
```python
import numpy as np

def forfact(n):
    result = 1
    for i in range(1, n+1):
        result *= i
    return result
```

<!-- grade_id: cell-21816c4d24477aee, points: 4, grade: true, solution: false -->
```python
n = 9
assert forfact(n) == np.prod(np.arange(1,n+1))
```

2. Write a function called **whilefact** that takes the input argument **n** and returns the factorial of n. I want you to solve it using a **while** loop

<!-- grade_id: cell-bc4d7b9f01f70741, points: 0, grade: false, solution: true -->
```python
import numpy as np

def whilefact(n):
    result = 1
    i = 1
    while i <= n:
        result *= i
        i += 1
    return result
```

<!-- grade_id: cell-7b4760b78688ef23, points: 4, grade: true, solution: false -->
```python
n = 10
assert whilefact(n) == np.prod(np.arange(1,n+1))
```

3. There are a number of approximations to the number $\pi$ that have been developed over the years.   
Leibniz's formula for approximating $\pi$ is

$$ \pi = 4 - 4/3 + 4/5 - 4/7 + 4/9 - 4/11 + 4/13 ....   $$

$$ \pi = \sum_{n=0}^{\infty}{(-1)^n\frac{4}{2n+1}} $$

Note two things to pay attention to about this sum. 
* The first term is equivalent to 4/1.  So, we are counting 1,3,5,7,9,11.... for the denominator.  That is where the 2*n+1 comes from 
* We are alternating between adding and subtracting numbers.  That is where the $(-1)^n$ comes from.    

Write a function **leibniz** sums the terms in Leibniz's formula for $\pi$ until the magnitude of the *change* in the sum is less than a user specified value **tol**. The function should return the approximation to $\pi$ and the number of terms (in that order)

<!-- grade_id: cell-a4c6bcc96def1e0a, points: 0, grade: false, solution: true -->
```python
import numpy as np

def leibniz(tol):
    pi_est = 0
    n = 0
    term = 4 / (2*n + 1)
    sign = 1

    while abs(term) >= tol:
        pi_est += sign * term
        n += 1
        sign *= -1
        term = 4 / (2*n + 1)

    return pi_est, n
        
```

<!-- grade_id: cell-c3e5effce4348037, points: 4, grade: true, solution: false -->
```python
tol = 0.01
pi,n = leibniz(tol)
assert n == 200 
assert np.isclose(pi,np.pi, rtol = 0.005)
```

4. The `input` function in python allows for keyboard input. 

For example:
         
         d = input('test input:')
         
* will prompt for an input with the text 'test input'
* this argument is called the **prompt** and will appear in a box below this cell.
* After you type your response (and Enter) the keyboard input in stored in a variable **d** 
* By default the variable is of type `str`. 

Write a function **get_response** which takes two input arguments
* string **prompt**
* list or array **validresponses**

The function **get_response** uses the function `input` and prompts the user for a variable **response** 
with the prompt given by **prompt**

If the keyboard input is one of the responses in **validresponses**, the function should print **response** and return the variable **response**  

If the keyboard input is not in the list or array **validresponses** the function should print the string, 

'Invalid Response, Try Again' 

and then prompt the subject for a response again. 

The function should keep prompting the user until the user gives a response in the list or array **validresponses**

*Hint 1*: First test how input works and examine the variables it stores.  Just type this simple example in the box below and run it:

    x = input('How does this look?:')
    print(x)

The input box appears below the code cell.  Find it, and enter different types of variables there and see that it always returns a string even if you enter a number.  

*Hint 2*: You need to keep prompting them until they give you a valid input.  This suggests a `while` statement to me, where you check if their input is valid and only stop asking for a response if they give you a valid input. 

*Hint 3*: To test if their keyboard input, stored in the variables **response** is one of the values in **validresponses** you can use list or array methods with the following logical statements. 

array:             
* np.any(validresponses == response)

list:              
* response in validresponses


<!-- grade_id: cell-81bd7770e6c3b142, points: 8, grade: true, solution: true -->
```python
def get_response(prompt, validresponses):
    reponse = input(prompt + " ")
    while response not in validresponses:
        print("Invalid Response, Try Again")
        response = input(prompt + " ")
    print(response)
    return response
```

```python
prompt = "Ready to code an experiment?"
validresponses = ['y','n']
response = get_response(prompt,validresponses)
```

```python

```